# CPN 321 Computer Tut 1 2017
## Solving Equations

## 1  Linear algebra (7)

Figure 1 shows a flow network. Assume
steady state operation. If there are two
components in the system and no energy balances are done, how many
variables and equations are there?  Calculate the degrees of freedom.

<img src="graph/massbalance.png">

How many specifications are there in the following paragraph?

  *Streams A and H are 5~kg/h and 10~kg/h of pure substance X and substance Y
  respectively.  Processes 2, 3 and 6 are splitters, each yielding two
  streams. The stream exiting each splitter on the right is twice the
  flowrate of the other exiting stream, but has the same composition.  Processes 1, 4 and 5 are
  pure mixers.*
 
 These specifications can be written in the following form for the mass balance and the component balance:


 $$A \mathbf{x} = \mathbf{b}$$


Find $A$ and $\mathbf{b}$ for each of these balances. Create one function named `solve_matrices` that uses `numpy.linalg.solve`. The function must take $A$ and $\mathbf{b}$ and return the solution to the equation above as a matrix. Use this function to solve for the flows as well as the composition (in terms of X) of the streams. 

The solution must be in the following form and order:

$$
\mathbf{x} = 
\begin{bmatrix}
B\\
C\\
D\\
E\\
F\\
G\\
H\\
I\\
J\\
K\\
\end{bmatrix}
$$

Store the solution matrices for the flowrates and compositions in variables named `s` and `x` respectively. 



<!-- grade_id: solving_equations_answer, points: 0, grade: false, solution: true -->
```python
# used Python 3 in Jupyter
#9 variables where A and H are known inputs, therefore 9 unknowns and 9 m.b. equations DoF = 0


import numpy
from scipy import linalg

A = numpy.matrix([[-2, 1, 0, 0, 0, 0, 0, 0, 0],
                  [ 0, 0, 1, 0, 0,-2, 0, 0, 0],
                  [ 0, 0, 0, 0, 0, 0, 0, -2, 1],
                  [ -1, -1, 0, 0, 1, 0, 0, 0, 0],
                  [ 0, 1, -1, 0, 0, -1, 0, 0, 0],
                  [ 0, 0, 0, 0, 0, 1, 0, -1, -1],
                  [ -1, 0, 0, 1, 0, 0, 0, 0, 0],
                  [ 0, 0, 0, -1, 0, 0, 1, 0, 0],
                  [0, 0, 0, 0, 1, 0, -1, -1, 0]])

b = numpy.matrix([[0, 0, 0, 0, 0, 0, 5, 10, 0]]).T
s = numpy.linalg.solve(A,b)
print("A=",5)
print("B=",s[0])
print("C=",s[1])
print("D=",s[2])
print("E=",s[3])
print("F=",s[4])
print("G=",s[5])
print("H=",10)
print("I=",s[6])
print("J=",s[7])
print("K=",s[8])

print (s)
                 
A=5
B=s[0,0]
C=s[1,0]
D=s[2,0]
E=s[3,0]
F=s[4,0]
G=s[5,0]
H=10
I=s[6,0]
J=s[7,0]
K=s[8,0]

B = numpy.matrix([[-B, 0, 0, E, 0, 0, 0, 0, 0],
                  [ 0, 0, 0, -E, 0, 0, I, 0, 0],
                  [ 0, 0, 0, 0, F, 0, -I, -J, 0],
                  [ -2*B, C, 0, 0, 0, 0, 0, 0, 0],
                  [ 0, -C, D, 0, 0, 0, 0, 0, 0],
                  [ 0, 0, 0, 0, 0, 0, 0, -2*J, K],
                  [ -B, -C, 0, 0, F, 0, 0, 0, 0],
                  [ 0, C, -D, 0, 0, -G, 0, 0, 0],
                  [0, 0, 0, 0, 0, G, 0, -J, -K]])

b = numpy.matrix([[5, 0, 0, 0, 0, 0, 0, 0, 0]]).T
x = numpy.linalg.solve(B,b)
print("xa=",1)
print("xb=",x[0])
print("xc=",x[1])
print("xd=",x[2])
print("xe=",x[3])
print("xf=",x[4])
print("xg=",x[5])
print("xh=",0)
print("xi=",x[6])
print("xj=",x[7])
print("xk=",x[8])

print (x)

print()       


#raise NotImplementedError()
```

Output:
```
A= 5
B= [[ 8.4375]]
C= [[ 16.875]]
D= [[ 11.25]]
E= [[ 13.4375]]
F= [[ 25.3125]]
G= [[ 5.625]]
H= 10
I= [[ 23.4375]]
J= [[ 1.875]]
K= [[ 3.75]]
[[  8.4375]
 [ 16.875 ]
 [ 11.25  ]
 [ 13.4375]
 [ 25.3125]
 [  5.625 ]
 [ 23.4375]
 [  1.875 ]
 [  3.75  ]]
xa= 1
xb= [[ 0.2962963]]
xc= [[ 0.2962963]]
xd= [[ 0.44444444]]
xe= [[ 0.55813953]]
xf= [[ 0.2962963]]
xg= [[-0.]]
xh= 0
xi= [[ 0.32]]
xj= [[-0.]]
xk= [[-0.]]
[[ 0.2962963 ]
 [ 0.2962963 ]
 [ 0.44444444]
 [ 0.55813953]
 [ 0.2962963 ]
 [-0.        ]
 [ 0.32      ]
 [-0.        ]
 [-0.        ]]
```

<!-- grade_id: solving_equations_tests, points: 7, grade: true, solution: false -->
```python
# DO NOT CHANGE CELL
# test for correct shape of s
assert numpy.shape(s) == (9, 1)
# test for correct shape of x
assert numpy.shape(x) == (9, 1)

```

```python

```

## 2 CIR 123 example 2.2 (6)

1000 kg/hour of raw material with the following analysis is available: $Na_2CO_3 - 80 \%; Na_2SO_4 - 15 \%; H_2O - 5\%$. It is required to produce 400 kg/hour $Na_2CO_3$ crystals (K) and as much $Na_2SO_4$ crystals (S) as possible and the following process is recommended:

  <img src="graph/process.png">
 

The raw material (G) is dissolved in water (W) to give a solution containing 28 % $Na_2CO_3$.
A part of the solution (V) is then cooled to 21 °C, which results in a part of the $Na_2CO_3$ crystallising as the decahydrate crystal ($Na_2CO_3.10H_2O$). These crystals are separated from the mother liquor. A crystallisation test showed that the mother liquor (M) contains 13.1 % $Na_2CO_3$ and 13.9 % $Na_2SO_4$. It can be assumed that the separation is perfect, in other words, the $Na_2CO_3.10H_2O$ stream (D) which enters evaporator (1) does not contain any mother liquor.

The water in the  $Na_2CO_3.10H_2O$ stream is evaporated in evaporator (1) to give pure $Na_2CO_3$ product (K).

The balance of the solution (N) which is not cooled, is mixed with the mother liquor (M) and the mixture is neutralised with $H_2SO_4$ according to the reaction


  $$Na_2CO_3 + H_2SO_4 \rightarrow Na_2SO_4 + H_2O + CO_2$$


The neutralised stream (R) is then evaporated in evaporator (2) to give $Na_2SO_4$ product (S).

In the original question, the following questions were asked:

* The mass flow rate of the 
$NaCO_3.10H_2O$ stream (D).
* The mass flow rate of the solution that is cooled (V).
* The mass flow rate of water (W) required in the solution tank.
* The mass $H_2SO_4$ solution required per hour if the concentration of the solution is 98 % $H_2SO_4$.
* The total mass of water evaporated per hour in evaporators (1) and (2).
* The mass $Na_2SO_4$ product (S) produced per hour.


Write down all the equations that would be used to calculate these unknowns, then use `sympy` in an IPython notebook to solve for them simultaneously. Store the equations in a variable named `equations`.


The variables have already been created for you, **DO NOT** change them. Your solution must contain a value for each of the variables. Store the solution in a variable named `solution`.

**HINT:** Use the solution you get from `sympy` as your final answer. 


<!-- grade_id: CIR123_answer, points: 0, grade: false, solution: true -->
```python
from sympy import *

G, W, N, V, P, M, H2SO4, CO2, R, H2Oe1, S, D, K, H2Oe2 = var('G, W, N, V, P, M, H2SO4, CO2, R, H2Oe1, S, D, K, H2Oe2')
unknowns = [G, W, N, V, P, M, H2SO4, CO2, R, H2Oe1, S, D, K, H2Oe2]  # Do not change the name of this variable.

MM_Na2CO3 = (2*23 + 12 + 3*16)
MM_Na2SO4 = (2*23 + 32 + 4*16)
MM_H2SO4 = (2 + 32 + 4*16)
MM_H20 = (2+16)
MM_CO2 = 12 + 2*16
# Basis: 1 hour
eq1 = G - 1000 

# The mass flow rate of the NaCO3.10H2O stream (D)
eq2 = K - 400 #output Na2CO3 required in Kg/hour

eq3 = D - ((K/MM_Na2CO3)*(MM_Na2CO3 + 10*MM_H20))

#The mass flow rate of the solution that is cooled (V).

eq4 = V - D - M  # mass balance around separator of crystals and mother liquor
eq5 = 0.28*V - 400 - 0.131*M #Na2CO3 balance around separator


#The mass flow rate of water (W) required in the solution tank.

eq6 = W + G - N - V #total mass balance
eq7 = 0.8*G - 0.28*(N + V) #Na2CO3 mass balance around dissolver

eq8 = P - N - M  #mass balance around mixer 

#The mass H2SO4 solution required per hour if the concentration of the solution is 98 % H2SO4.

eq9 = H2SO4 - ((0.28*N + 0.131*M)/MM_Na2CO3)*MM_H2SO4/0.98   

#The total mass of water evaporated per hour in evaporators (1) and (2).
#The mass Na2SO4N product (S) produced per hour.
# Na2CO3+H2SO4→Na2SO4+H2O+CO2 in neutraliser

eq10 = S - 0.15*G - ((0.28*N + 0.131*M)/MM_Na2CO3)*MM_Na2SO4

eq11 = H2Oe1 + K - D #mass balance around Evaporator1
eq12 = CO2 - ((0.28*N + 0.131*M)/MM_Na2CO3)*MM_CO2 
eq13 = G + W + H2SO4 - H2Oe1 - H2Oe2 - K - S - CO2

eq14 = R + CO2 - H2SO4 - P # mass balance around neutraliser


equations = [eq1, eq2, eq3, eq4, eq5, eq6, eq7, eq8, eq9, eq10, eq11, eq12, eq13, eq14]

solve(equations, [G, W, N, V, P, M, H2SO4, CO2, R, H2Oe1, S, D, K, H2Oe2])

solution =(G, W, N, V, P, M, H2SO4, CO2, R, H2Oe1, S, D, K, H2Oe2)


raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-96-3f5e2b675752> in <module>()
     54 
     55 
---> 56 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: CIR123_answer_check, points: 6, grade: true, solution: false -->
```python
# DO NOT CHANGE CELL
# test for correct number of equations
assert len(equations) == 14

# test to see if you used the unknowns as given
a = set()
for e in equations:
    a = a.union(e.free_symbols)

assert numpy.asarray(unknowns).all() in a

# test to check number of solutions
assert len(solution) == 14


```

## 3 Root finding (12)

Compare the efficiency of different algorithms in solving for $f'$ in
the following nonlinear equation (the Colebrook estimator for the
friction factor).
\begin{equation}
  \frac{1}{\sqrt{f'}} = -2\log \left ( \frac{\epsilon/D}{3.7} +
    \frac{2.51}{Re\sqrt{f'} }\right )
\end{equation}

Assume that the flow is water in a 1/4" pipe at 280 ft$^3$/hour with
$\epsilon=0.007$ mm.
 
Generate a plot and indicate on the plot where a possible solution might be. Remember to add axis labels and a legend to you plot.

Write you code in the cell below.


<!-- grade_id: root_finding_answer_part1, points: 1, grade: true, solution: true -->
```python
import numpy
import matplotlib.pyplot as plt
%matplotlib inline
# YOUR CODE HERE
raise NotImplementedError()


```

Create a script that will plot a convergence graph for the following algorithms using **50** function evaluations. The graph will show the absolute value of the difference between the LHS and RHS on the y axis and the number of function evaluations on the x axis. On each graph also plot the results of the previous algorithms, i.e. the graph for Newton-Raphson will show the results for both Fixed point iteration and Newton-Raphson. The last graph must show the results of all the algorithms. 


### 3.1 Fixed point iteration
Isolate $f'$ in implicit form from the
  equation and then use the result repeatedly.  Note that there are
  two forms for this and typically only one will be convergent.  Calculate the error for both. Start at $f'=0.02$.
  
  For information on fixed point iteration click [here](https://en.wikipedia.org/wiki/Fixed-point_iteration)



Write you code in the cell below.

**HINT:** You can repeat the `matplotlib` plot statement on a new line to add another dataset to the graph.

<!-- grade_id: root_finding_answer_part_2, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### 3.2 Newton-Raphson
  Use the same starting point as above (remember, you need the derivative).
  
   For information on Newton-Raphson click [here](https://en.wikipedia.org/wiki/Newton%27s_method)
  
  
  Write you code in the cell below.


<!-- grade_id: root_finding_newton_raphson, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

 ### 3.3 Regula-Fals (False point)
 Assume the interval is
  between $f'=0.001$ and $f'=0.5$.
  
  For information on Regula-falsi click [here](https://en.wikipedia.org/wiki/False_position_method)
  
  Write you code in the cell below


<!-- grade_id: root_finding_regula_falsi, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

 ### 3.4 Interval halving
 Assume the same interval as above.

 For information on interval-halving click [here](https://en.wikipedia.org/wiki/Bisection_method)

Write you code in the cell below.

<!-- grade_id: root_finding_interval_halving, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## 4 Integration (5)
A tank with one inflow and one outflow has an outflow-height
dependency described by

\begin{equation}
  F_{out} = K \times \sqrt{h}
\end{equation}

where $K$ is a suitable constant.

The dynamic mass balance simplifies with the assumption of constant $\rho$ to
\begin{equation}
  \frac{dh}{dt} = \frac{F_{in} - F_{out}}{A}
\end{equation}

where $A$ is the cross-sectional area of the tank.

Investigate the effect of stepsize on the accuracy of Euler integration by comparing to `scipy.integrate.odeint`.

Call your Euler function `euler`, use the following function definition: `def euler(dxdt, x0, tspan):`.
Create one function describing the system called
`system` in the form required by `odeint` and write your
integration routine so that it functions similarly to the Python integration routine.

Use $K=1.2$, $A=1$.
The initial height $h(0)=2$.
Simulate an immediate change in the input flowrate to $F_{in}(t)=2.5$ for $t>0$.

Use a large step size and a small step size to see the difference between the algorithms.

Define your functions in the cell below.


<!-- grade_id: integration_functions_create, points: 3, grade: true, solution: true -->
```python
import numpy
import matplotlib.pyplot as plt
import scipy.integrate

# YOUR CODE HERE
raise NotImplementedError()
```

 Plot the results of both algrithms on one graph. Be sure to include both large and small step sizes for euler.
 
 Remember to add axis labels and a legend to you plot.
 
 Define your plotting code in the cell below.

<!-- grade_id: integration_comparison_plot, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: integration_functions_test, points: 1, grade: true, solution: false -->
```python
# DO NOT CHANGE CELL
```