## 1 Deadtime (10)

Develop a simulation of the system in Example 5.3, simulating the same situation as in the solution presented. 

### 1.1 Numerical
Simulate the system with deadtime numerically using `numpy` for a value of $2000 \frac{kg}{m^3}$ for $c_i$. For the dead time, you can interpolate on the values from the past using `numpy.interp`. Define a function called `numeric` that takes the timespan as an input and returns the numerical response of $c_1'$.




<!-- grade_id: cell-5a309567e9a2649c, points: 0, grade: false, solution: true -->
```python
import numpy
import matplotlib.pyplot as plt
# Do not change this, %matplotlib notebook makes it difficult to mark
%matplotlib inline 


ci = 2000  # kg/m^3
ts = numpy.linspace(0, 100, 1000)  # timespan DO NOT CHANGE


def numeric(ts):
    
    return

# YOUR CODE HERE
```

<!-- grade_id: cell-64c1f71e45cd70d9, points: 2, grade: true, solution: false -->
```python
# DO NOT CHANGE
assert numpy.round(numeric(ts)[10],2) == 2.02

```

Plot the response below.

<!-- grade_id: cell-c4d199dd27937f9a, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

### 1.2 Deadtime approximation
Use `sympy` to calculate an approximation for the deadtime. Create a function called `dead_time_approx` that takes the timespan as input. This function must return the resulting $c_1'$ when using an approximation for the dead time.

NOTE: Do not round any values by hand. Rather use the variables in your equations. The reason for this is that `sympy` calculates values very precisely, if you round values you will not get the correct answer and the grader will fail you on this question.

<!-- grade_id: cell-4578e9efb2af74ae, points: 0, grade: false, solution: true -->
```python
import sympy

def dead_time_approx(ts):
    
    return

# YOUR CODE HERE
```

<!-- grade_id: cell-bc2b58b07e16db0d, points: 3, grade: true, solution: false -->
```python
# DO NOT CHANGE
assert numpy.round(list(map(float, dead_time_approx([1])))) == [2]
```

Plot the response below.

<!-- grade_id: cell-c48fc3c761d9e245, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

### 1.3 Analytical solution
Create a function that returns the time domain response of **Eq 5-56** in the textbook. Name this function *analytical*, it must take the time span as an input. **NOTE:** you will need to use `numpy.interp` for this solution

<!-- grade_id: cell-07f7409b5c944f31, points: 0, grade: false, solution: true -->
```python
def analytical(ts):
    
    return
# YOUR CODE HERE
```

<!-- grade_id: cell-068c05f922a82ea6, points: 2, grade: true, solution: false -->
```python
# DO NOT CHANGE
assert round(analytical(ts)[100], 2) == 124.45

```

Plot the response of Eq 5-56 below.

<!-- grade_id: cell-db789eef2fd116b4, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

## 2 Data reading (8)
Next, modify your program to read the value of $c_i$ from the attached file `data.csv` and simulate the system. You can use the `pandas` module to read the file. Write a function that uses `pandas` to read the csv data. Name the function `read_data` it should take the file name as an input and return the data. Click [here](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.read_csv.html) to see the documentation on how to read a csv with `pandas`. `pandas` creates a DataFrame object that works like a dictionary, i.e. it has key value pairs. Remember to change the timestamps into delta time in terms of minutes. This conversion must also be done in your `read_data` function and added to the DataFrame. 


For a quick tutorial on `pandas` click [here](https://pandas.pydata.org/pandas-docs/stable/10min.html).


Write the code for reading the data below.

<!-- grade_id: cell-d3ee2dff5d5d3ab2, points: 0, grade: false, solution: true -->
```python
import numpy
import matplotlib.pyplot as plt
import pandas as pd


def read_data(filename):
    
    return

# YOUR CODE HERE
```

<!-- grade_id: cell-108261f1392016fb, points: 2, grade: true, solution: false -->
```python
# DO NOT CHANGE
```

Write a function called `simulate` that takes the data from the `read_data` function and the timespan as inputs and returns the simulated values for $c_1'$. Use the numerical method to simulate the system, question 1.1. Use 10000 points in your timespan. You will need to use `numpy.interp` to interpolate between the values in the data file.

<!-- grade_id: cell-946be71c5d4154b3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-5c4279f09519187d, points: 4, grade: true, solution: false -->
```python
# DO NOT CHANGE
assert len(simulate(data, ts)) == 10000
assert round(simulate(data, ts)[-1]) == 182

```

Using subplots plot the response of the system ($c_1'$) and the changing values of $c_i$ below with respect to time.

<!-- grade_id: cell-2d190ec9a9577892, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

```python

```