<center><h2>CPN 321</h2></center>
<center><h3>C4 2017</h3></center>

In this tut we will be analysing data and comparing it to theoretical predictions.
"LPG Sample dataset" contains a P&ID describing a unit on a large chemical plant, consisting of two distillation columns and ancilliary equipment. The CSV data files contain values for the various controllers on the plant sampled once a minute.
In the datafile, each controller has three values associated with it.


* Process value (PV): the measurement from the plant - typically the controlled variable

* Output (OP): the output of the controller. This may be a valve fraction if the controller is connected to a valve, or a setpoint for another controller.

* Setpoint (SP): target value for the controller

The values in the files have been normalised by subtracting the minimum and dividing by the range.

Use the first data file `LPG_Data_Set_1_n.csv`.

## 1 Non-dynamic regression (14)

You can open the files in Excel by double-clicking them, in Python by using the `pandas.read_csv` function.

Find the relationship between the valve fraction `F1023.OP` and the flowrate out of `DM108 (F1023.PV)`, neglecting dynamics. Try this in Excel and Python. Call your Excel spreadsheet `Regression.xlsx`. Have a look at [this notebook](http://nbviewer.jupyter.org/github/alchemyst/Dynamics-and-Control/blob/master/Regression.ipynb) to see how regression works in Python


### 1.1 Read the data

Write a function called `read_csv` that uses pandas to read the csv data. The function must take the filename as an input.

<!-- grade_id: cell-c29386c6f9030d78, points: 0, grade: false, solution: true -->
```python
import numpy
import matplotlib.pyplot as plt
import scipy.optimize
import scipy
import pandas
%matplotlib inline

def read_data(filename):
    
    return

# YOUR CODE HERE
```

<!-- grade_id: cell-273ea21179eaa5f4, points: 2, grade: true, solution: false -->
```python
assert read_data('LPG_Data_Set_1_n.csv').l1013aspv.iloc[0] == 0.41894060995184573
```

Plot the *f1023aspv* (x-axis) and the *f1023aspv* (y-axis) against each other below as a scatter plot. Use an alpha of 0.05 for the markers.

<!-- grade_id: cell-3f0049ee18c201df, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

### 1.2 Linear regression
It is now time to fit a model to the data. Write a function called `linear_fit` that takes the x and y values a inputs and returns the model coefficients as a list. Use the `numpy.ployfit` function to fit a model to the data. For info on how this function works click [here](https://docs.scipy.org/doc/numpy/reference/generated/numpy.polyfit.html).

<!-- grade_id: cell-8e53f6fa1147674f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE



```

<!-- grade_id: cell-6eacdf8f525f8d87, points: 2, grade: true, solution: false -->
```python
data = read_data('LPG_Data_Set_1_n.csv')
assert round(linear_fit(data.f1023asop, data.f1023aspv)[0], 2) == 1.03
```

Below plot the raw data and the linear model on one graph. Remeber to add a ledgend to the plot. Make the color of the first order line red.

<!-- grade_id: cell-cd066a98e0f05134, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

### 1.2 Non-linear regression
Now fit the following nonlinear relationship on the data.  


$$F=a\tanh(x/b-c)+d$$ 


You can to this by minimising the sum of the squares of the error. The `scipy.optimize.curve_fit` function will be useful. Fore more info on how `scipy.optimize.curve_fit` works click [here](https://docs.scipy.org/doc/scipy/reference/generated/scipy.optimize.curve_fit.html).  

Write a function named `nonlinear_function` that takes x, a, b, c and d as inputs and returns the function value at those specified values.

<!-- grade_id: cell-45ea9653279dc325, points: 0, grade: false, solution: true -->
```python
def nonlinear_function(x, a, b, c, d):
    
    return 

# YOUR CODE HERE
```

<!-- grade_id: cell-b3e3479376e041b1, points: 1, grade: true, solution: false -->
```python
assert round(nonlinear_function(2, 1, 1, 1, 1), 2) == 1.76
```

Write a function called `nonlinear_fit` that takes the pandas DataFrame as input and returns the parameters for the non-linear fit.

<!-- grade_id: cell-99e6f482728b3bb5, points: 0, grade: false, solution: true -->
```python
def nonlinear_fit(data):

    return 

# YOUR CODE HERE
```

<!-- grade_id: cell-9df91a9381217569, points: 3, grade: true, solution: false -->
```python
assert round(nonlinear_fit(data)[0], 2) == 1.03 
```

Below plot the non-linear fit along with the raw data and linear fit solution on one graph. Remember to give your graph a ledgend.

<!-- grade_id: cell-46c3100f0390b655, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

## 2 Dynamic models (13)
### 2.1 Simulation

The linear time invariant state space format can be written as

$$\dot x=Ax+Bu$$
$$y=Cx+Du$$

<br />
For a second order system the linear time invariant state space representation is
<br />
$$G(s) = \frac{y(s)}{u(s)} = \frac{K}{\tau^2s^2 + 2\tau\zeta s + 1}$$

$$ A = \begin{bmatrix} 0 & 1\\ -\frac{1}{\tau^{2}} & - \frac{2 \zeta}{\tau} \end{bmatrix}  B =  \begin{bmatrix} 0\\ \frac{K}{\tau^{2}} \end{bmatrix}$$


$$C= \begin{bmatrix} 1& 0 \end{bmatrix} D= \begin{bmatrix} 0 \end{bmatrix} $$


The state space representation can be used to integrate (using Euler) to find the time response

### 2.2 Fitting
Use the given simulation methods to find the parameters of an overdamped second order transfer function relating the level in DM108 (`L1013.PV`) to the flow out of that drum (`F1023.PV`). These tags are in the same datafile as the one used for the above questions. 

First use an estimate and manual tuning to plot the response. 

Simulate a state space model using Euler integration. Refer to [this link](http://nbviewer.jupyter.org/github/alchemyst/Dynamics-and-Control/blob/master/Simulation%20of%20arbitrary%20transfer%20functions.ipynb) on how to deal with Linear, Time Invariant (LTI) systems 

Now repeat the fit finding the parameters of equation 6-36 using `numpy.linalg.lstsq`

Produce a graph showing the fits.

Hint: in order to see the relationships more clearly, you can plot only the first thousand or so data points.

Call your file `secondorderfit.py`


Write a function called `norm_data` that takes the pandas DataFrame as input and return the normalized data as output, i.e. your function must return two list each consisting of the normalized values of the required columns. The order in which they must be return is specified below.

<!-- grade_id: cell-9a478cce630000c7, points: 0, grade: false, solution: true -->
```python
import numpy
import matplotlib.pyplot as plt
%matplotlib inline

finalT = 1000
# Data is sampled at one minute intervals
t = numpy.arange(finalT)

def norm_data(data):
    unorm = []  # you kust change this
    ynorm = []  # you kust change this
    
    return unorm, ynorm

# YOUR CODE HERE
```

<!-- grade_id: cell-579d1227323d235b, points: 1, grade: true, solution: false -->
```python
filename = 'LPG_Data_Set_1_n.csv'
data = read_data(filename)
assert len(norm_data(data)[0]) == 1000
assert round(norm_data(data)[0][0], 2) == 0.71
assert round(norm_data(data)[1][0], 2) == 0.88
```

Plot the normalized data below using subplots. Remeber to add labels to your data.

<!-- grade_id: cell-113ae647ead948b0, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

### 2.3 Linear, Time Invariant system

Create a function called `state_space` that takes $\tau_1$, $\tau_2$ and $K$ as input and returns the state space matrices $A$, $B$, $C$ and $D$. You can use `scipy.signal.lti()` to construct the transfer function and get the state space.

<!-- grade_id: cell-0a2ed0408ae27baf, points: 0, grade: false, solution: true -->
```python
# Determine values for K, tau1 and tau2 by simulating the second order response
def state_space(K, tau1, tau2):
    
    return

# YOUR CODE HERE
```

<!-- grade_id: cell-b1e80ab907c7dbbf, points: 2, grade: true, solution: false -->
```python
K = -1.1
tau1 = 80
tau2 = 30
A, B, C, D = state_space(K, tau1, tau2)
numpy.round(A, 6)
assert (numpy.round(A, 6) == numpy.array([[ -4.58330000e-02,  -4.17000000e-04],[  1.00,   0.00]])).all()
```

### 2.4 Nonlinear fit
Write a function called `y_predict` that takes the input data ($u$), $K$, $\tau_1$, $\tau_2$ and the time as inputs and predicts the output ($y$). You can use `scipy.signal.lsim` to simulate the get the $y$ data from the LTI and $u$ data. Click [here](https://docs.scipy.org/doc/scipy-0.16.1/reference/generated/scipy.signal.lsim.html) for more info. Use the same time variabale as in 2.2.

<!-- grade_id: cell-1001d8070a6947ec, points: 0, grade: false, solution: true -->
```python
def y_predict(udata, K, tau1, tau2):

    return ypred

# YOUR CODE HERE
```

<!-- grade_id: cell-6172f8346df054ef, points: 2, grade: true, solution: false -->
```python
filename = 'LPG_Data_Set_1_n.csv'
K = -1.1
tau1 = 80
tau2 = 30
data = pandas.read_csv(filename)
udata = data.f1023aspv[:finalT]
ydata = data.l1013aspv[:finalT]
assert round(y_predict(udata, K, tau1, tau2)[-1], 2) == 0.3

```

Create a function called `raw_nonlinear_fit` that uses `scipy.optimize.curve_fit` and your `y_predict` function to the second order system to the data and returns the fitted values for $K$, $\tau_1$ and $\tau_2$.

<!-- grade_id: cell-91a021bbc2714a38, points: 0, grade: false, solution: true -->
```python
beta = [K, tau1, tau2]

def raw_nonlinear_fit():
    
    return 

# YOUR CODE HERE
```

<!-- grade_id: cell-a5a8bfffce5d926b, points: 1, grade: true, solution: false -->
```python
assert round(raw_nonlinear_fit()[-1], 2) == 70.54

```

### 2.5 Linear fit

Use equation 6-36 to fit a linear model to the raw data. Create a function called `discrete_coeffs` that returns the coefficients needed to calculate equation 6-36 in the following order $a_1$, $a_2$, $b_1$, $b_2$. Have a look at example 6.5 in the handbook. You can use `scipy.linalg.toeplitz` function (click [here](https://docs.scipy.org/doc/scipy-0.14.0/reference/generated/scipy.linalg.toeplitz.html)) and the `numpy.linalg.lstsq` function (click [here](https://docs.scipy.org/doc/numpy/reference/generated/numpy.linalg.lstsq.html)).

<!-- grade_id: cell-f0c32b76b6e54d0b, points: 0, grade: false, solution: true -->
```python
def discrete_coeffs(udata, ydata):
    
    return


# YOUR CODE HERE
```

<!-- grade_id: cell-9c34f3c12ebebc18, points: 3, grade: true, solution: false -->
```python
a1, a2, b1, b2 = discrete_coeffs(udata, ydata)
assert round(a1, 2) == 0.57
```

Use the function below to simulate the linear fitted function.

```python
def discretesim(t, a, b, u):
    """ Simulate discrete system on arbitrary input
        Currently only supports second order systems
    """
    a1, a2 = a
    b1, b2 = b
    ys = []
    y = y_1 = y_2 = 0
    for i, ti in enumerate(t):
        u_1 = 0 if i < 1 else u[i-1]
        u_2 = 0 if i < 2 else u[i-2]
        y_2, y_1 = y_1, y

        # Eq 6-36
        y = a1*y_1 + a2*y_2 + b1*u_1 + b2*u_2
        ys.append(y)

    return numpy.array(ys)
```

Plot figures of linear and non-linear fittings

<!-- grade_id: cell-9667eac300f767e5, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

```python

```